```python
# Initialize Otter
import otter
grader = otter.Notebook("mod01.ipynb")
```

# Module 1: Intro to Human Geography!

### Welcome to `GEOG 102`

In this interactive notebook, we will work as geographical data scientists who would like to answer questions, learn and/or make predictions by looking at data. The first step in this quest is to look at collected data, visualize the data and observe. This will bring up all kinds of questions.

This activity's goal is to review some of the material presented in the videos and lectures of module 1 by:
- Importing data collected by the World Trade Organization (WTO) and the Food and Agriculture Organization (FAO) of the United Nations
- Focusing in on data of interest by looking at it in a tabular format
- Visualizing the data using graphs, barcharts and maps
- Observe, ask questions and learn from the data

We will ...
- Introduce jupyter notebooks as an interactive learning platform
- Review some concepts and terminology around maps
- Compare 2 countries (USA and Haiti), their location, colonial history & uneven economic regional growth

TODO: review the section above with input from Geordie.



# 1. What is a Jupyter notebook?
This webpage is called a Jupyter notebook. A notebook is a place to write code and view their results, and also to write text.
In a notebook, each rectangle containing text or code is called a *cell*.

**Text cells** (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.
After you edit a text cell, click the "run cell" button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. 

**Code cells** contain code in the Python 3 language. Running a code cell will execute all of the code it contains.
To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶| or hold down `shift` + `return`.

**Activity 1:** This is a text cell. It is the cell type where we can type text that isn't code. Go ahead and double click in this cell and you will see that you can edit it. 

**Type something here:** ....


**Activity 2:** Click on the code cell below and run the code:

```python
#This cell is a code cell. It is where we can type code that can be executed.
#The hashtag at the start of this line makes it so that this text is a comment not code. 

print("Hello, World! \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

And this one:

```python
# This coding cell imports some python libraries that we will be using throughout this notebook
# Don't worry about what they are, just run this cell before running any other cells below this one
from datascience import *
import numpy as np
import otter
grader = otter.Notebook("mod01.ipynb")

%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

import folium
def createMap(locationInit, zoomInit, marker1Location, marker1Name, marker2Location, marker2Name):
    myMap = folium.Map(location=locationInit, tiles="OpenStreetMap", zoom_start= zoomInit)
    folium.Marker(marker1Location, popup = marker1Name).add_to(myMap)
    folium.Marker(location = marker2Location, popup = marker2Name).add_to(myMap)
    return myMap

from IPython.display import YouTubeVideo
```

# 2. Comparing 2 locations and their colonial history
Before you continue, make sure to watch the following 2 lectures posted on canvas:
- Regional economic growth lecture! 
- Uneven economic development lecture: Haiti

# 2.1 Maps
TODO: Add image here
In this section we will look at the map, the 2 first colonies in the Americas that become independent and their location. We will compare the difference in their geographies in terms of latitude, climate and discuss how this affected who settled there. 
In this section, we will look at how the location of the 2 first colonies in the Americas that become independent and their location: **The 13 colonies (early USA)** and **Haiti**. 

Both locations have similarities and differences when locating them on the map:

**Question 2.1** Let's first have a look at their locations on the map! In the coding cell below:
- Set **penn_lat** and **penn_long** equal to the latitude and longitude of Philadelphia, PA (the original capital of the 13 colonies)
- Set **port_au_prince_lat** and **port_au_prince_long** equal to the latitude and longitude of Port-au-Prince, Haiti (the capital of Haiti)

HINT: You can use [Google Maps](https://www.google.com/maps) and right click on a location to get the coordinates (latitude and longitude).


```python
#replace the ... with the correct latitudes and longitudes
penn_lat = ...
penn_long = ...
port_au_prince_lat = ...
port_au_prince_long = ...
createMap([33,-100], 2, [port_au_prince_lat, port_au_prince_long], 'Port-au-Prince (Haiti)', [penn_lat, penn_long], 'Philadelphia, PA (13 Colonies)')
```

```python
grader.check("q21")
```

**Question 2.2** Look at the map above. Which of the following statements is correct? 

Port-au-Prince (Haiti) and Philadephia, PA (USA) have...
1. similar latitudes and similar longitudes.
2. similar latitudes, but different longitudes.
3. different latitudes, but similar longitudes.
4. different latitudes and different longitudes.

Answer in the code cell below, by setting the name (variable) **correct** equal to 1, 2, 3 or 4.

```python
correct = ...
```

```python
grader.check("q22")
```

Run the code cell below to watch the video "Why all world maps are wrong" and answer question 2.3.

```python
# How accurate is the weather forecast?
# The original URL is: 
#https://youtu.be/kIID5FDi2JQ
YouTubeVideo("kIID5FDi2JQ")
```

**Question 2.3:** What do geographers call the projection of the 2D map created after question 2.1?
1. Mercator Projection
2. True Projection
3. Gall-Peters Projection
4. Winkel Tripel Projection
Answer in the code cell below, by setting the name (variable) **projection** equal to 1, 2, 3 or 4.

```python
projection = ...
```

```python
grader.check("q23")
```

# 2.2 Uneven Regional Economic Growth - USA versus Haiti

In lecture, we have discussed the story of Haiti and the uneven economic development during and after colonization. 

## 2.2.1 History: What too few textbooks told you

Before we look at some data. Let's run the code cell below to watch a TEDEd talk on the atlantic slave trade. This video will help with the next question 2.4.

```python
# The original URL is: 
# https://youtu.be/3NXC4Q_4JVg

YouTubeVideo("3NXC4Q_4JVg")
```

TODO: Add a question about the video here

**Question 2.4:** This is the question

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Now let's get started with some **data exploration**!

## 2.2.1 Merchandise Export over the years

We will focus on the economies of the USA and Haiti. The datasets used are from the World Trade Organization (WTO): https://stats.wto.org/.

```python
# Data comparing US versus haiti merchandise export between 1948 and 2023: Exports in million US $ 
# TODO: is this data useable for a question??
exports_data = Table.read_table('WtoData_20240516_export.csv')
exports_data
```

TODO: Work with Geordie to come up with a good question based on this data... Can this data be used??

```python
# Haiti merchandise export in millions of $
exports_data.plot('Year', 'Haiti')
Haiti_export_1948 = exports_data.column('Haiti').item(0)
Haiti_export_2023 = exports_data.column('Haiti').item(75)
growth_Haiti = round(Haiti_export_2023/Haiti_export_1948, 2)
print('Export growth in millions of $ (1948 versus 2023):\n')
print('Haiti increased from '+ str(Haiti_export_1948) + ' to ' + 
      str(Haiti_export_2023)+' million $, growth factor = ' + str(growth_Haiti)+'\n')
# USA merchandise export in millions of $
exports_data.plot('Year', 'USA')
USA_export_1948 = exports_data.column('USA').item(0)
USA_export_2023 = exports_data.column('USA').item(75)
growth_USA = round(USA_export_2023/USA_export_1948, 2)
print('USA increased from '+ str(USA_export_1948) + ' to ' + 
      str(USA_export_2023)+' million $, growth factor = ' + str(growth_USA))
```

TODO: Discuss this with Geordie and add a question

**Question 2.5:**
Maybe we could ask students how the plots look different? Are there any decades in particular that show a difference? Why would this be?

-> The plot for Haiti has a lot more noise (lots of ups and downs) compared to the plot for the US. This may imply an unstable economy.
Notice Haiti's merchandise export numbers did not grow before 1970, followed by a some growth followed by a drastic dip in the 1990s. What happened in Haiti in the 1990s that affected Haiti's export numbers? The merchandise export did not start growing again until 2000.

```python
test1 = ...
```

```python
grader.check("q25")
```

## 2.2.3 The cost and affordability of a healthy diet

Now let's look at the cost and affordability of a healthy diet in Haiti versus the US. TODO: Should we continue with a focus on Haiti or should we look at regions of the world or both??

The following data comes the Food and Agriculture Organization of the United Nations https://www.fao.org/faostat/en/#data.

```python
#TODO: Should we add this video? Or another one? How to feed the world video:
# The original URL is: 
# https://youtu.be/QloMOOG-bbE

YouTubeVideo("QloMOOG-bbE")
```

```python
# Use site: https://www.fao.org/faostat/en/#data
CAHD_data_USvsHaiti = Table.read_table('FAOSTAT_HealthyDiet_USvsHaiti.csv')
CAHD_data_USvsHaiti = CAHD_data_USvsHaiti.select('Area', 'Item', 'Year', 'Unit','Value')
CAHD_data_USvsHaiti.show()
```

We will focus in on 2021. Run the cell below to compare the cost of a health diet per person per day in the USA versus Haiti.
# TODO: Maybe we could compare 2017 with 2021? What do we see?

```python
# What is the cost of a healthy diet per person per day in the US versus Haiti in 2021?
CAHD_data_USvsHaiti.where('Item','Cost of a healthy diet (PPP dollar per person per day)').where('Year', 2021)
```

Possible question: looking at the data above, is the cost for a healthy diet the same in the US as it is in Haiti? Compare the numbers for 2021.

```python
pop_percent_healthy = CAHD_data_USvsHaiti.where('Item','Percentage of the population unable to afford a healthy diet (percent)').select('Area','Year','Value')
pop_percent_healthy.show()
```

```python
pop_percent_avg = pop_percent_healthy.group('Area',np.average).select('Area', 'Value average').relabeled('Value average', 'percent unable to afford healthy diet')
pop_percent_avg
```

```python
# Looking at data from 2021
pop_starch = CAHD_data_USvsHaiti.where('Item','Cost of starchy staples').select('Area','Value').relabeled('Value', 'Cost Starch staples') 
pop_animalsource = CAHD_data_USvsHaiti.where('Item','Cost of animal source foods').select('Area','Value').relabeled('Value', 'Cost Animal Source')
pop_leg_nuts_seeds = CAHD_data_USvsHaiti.where('Item','Cost of legumes, nuts and seeds').select('Area','Value').relabeled('Value', 'Cost legumes/nuts/seeds')
pop_vegetables = CAHD_data_USvsHaiti.where('Item','Cost of vegetables').select('Area','Value').relabeled('Value', 'Cost vegetables')
pop_fruits = CAHD_data_USvsHaiti.where('Item','Cost of fruits').select('Area','Value').relabeled('Value', 'Cost fruits')
pop_oils = CAHD_data_USvsHaiti.where('Item','Cost of oils and fats').select('Area','Value').relabeled('Value', 'Cost of oils and fats')
pop_cost_foods = pop_starch.join('Area', pop_animalsource,'Area').join('Area', pop_leg_nuts_seeds, 'Area').join(
    'Area', pop_vegetables, 'Area').join('Area', pop_fruits, 'Area').join('Area', pop_oils, 'Area')
# Cost of food per category
pop_cost_foods
```

```python
pop_cost_foods.barh('Area')
title_cost = plt.title('Cost in \$ per person per day in 2021')
```

**Question :** TODO add question

```python
todo = ...
```

```python
grader.check("q26")
```

### Should we add section below or just focus on Haiti versus US?

Taking a step back. Looking at the different continents.


TODO: discuss with Geordie if we should include the part below on the different continents (maybe better to just focus on comparing US and Haiti?)

```python
CAHD_data_regions = Table.read_table('FAOSTAT_HealthyDiet_regions.csv')
CAHD_data_regions = CAHD_data_regions.select('Area', 'Item', 'Year', 'Unit','Value')
CAHD_data_regions
```

```python
pop_percent_healthy_regions = CAHD_data_regions.where('Item','Percentage of the population unable to afford a healthy diet (percent)').select('Area','Year','Value')
pop_percent_healthy_regions.show()
```

```python
pop_percent_avg_regions = pop_percent_healthy_regions.group('Area',np.average).select('Area', 'Value average').relabeled('Value average', 'percent unable to afford healthy diet')
pop_percent_avg_regions
```

```python
pop_percent_avg_regions.barh('Area')
```

```python
pop_starch = CAHD_data_regions.where('Item','Cost of starchy staples').select('Area','Value').relabeled('Value', 'Cost Starch staples') # only data available for 2017
pop_animalsource = CAHD_data_regions.where('Item','Cost of animal source foods').select('Area','Value').relabeled('Value', 'Cost Animal Source') # only data available for 2017
pop_leg_nuts_seeds = CAHD_data_regions.where('Item','Cost of legumes, nuts and seeds').select('Area','Value').relabeled('Value', 'Cost legumes/nuts/seeds') # only data available for 2017
pop_vegetables = CAHD_data_regions.where('Item','Cost of vegetables').select('Area','Value').relabeled('Value', 'Cost vegetables') # only data available for 2017
pop_fruits = CAHD_data_regions.where('Item','Cost of fruits').select('Area','Value').relabeled('Value', 'Cost fruits') # only data available for 2017
pop_oils = CAHD_data_regions.where('Item','Cost of oils and fats').select('Area','Value').relabeled('Value', 'Cost of oils and fats') # only data available for 2017
pop_cost_foods = pop_starch.join('Area', pop_animalsource,'Area').join('Area', pop_leg_nuts_seeds, 'Area').join('Area', pop_vegetables, 'Area').join('Area', pop_fruits, 'Area').join('Area', pop_oils, 'Area')
pop_cost_foods
# cost unit: PPP $ per person per day
healthy_diet_summary_regions = pop_percent_avg_regions.join('Area', pop_cost_foods, 'Area')
healthy_diet_summary_regions

```

```python
healthy_diet_summary_regions.drop('percent unable to afford healthy diet').barh('Area')
```

**CONGRATULATIONS!** You just finished your first jupyter notebook assignment! 

We hope you enjoyed seeing the power of data and computer programming, tools that can help us understand how geography affects humans around the world!

Be sure to...

- run all of the tests and verify that they all pass, 
- choose **Download as PDF** from the **File** menu
- submit the .pdf file on **canvas**.

Tests for q21 (points: None)
```
>>> 39 <= penn_lat <= 40
True
```
```
>>> -76 <= penn_long <= -75
True
```
```
>>> 18 <= port_au_prince_lat <= 19
True
```
```
>>> -73 <= port_au_prince_long <= -72
True
```


Tests for q22 (points: None)
```
>>> correct == 3
True
```


Tests for q23 (points: None)
```
>>> projection == 1
True
```


Tests for q25 (points: None)
```
>>> test1 == 0
True
```


Tests for q26 (points: None)
```
>>> todo == 1
True
```
