```python
# Initialize Otter
import otter
grader = otter.Notebook("mod02.ipynb")
```

# Human Geography: Development Geographies



```python
# Don't change this cell; just run it. If you need additional libraries for your project, you can add them to this cell.

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets

```

In the following cells we will look at the different socioeconomic and sociodemographic variables used to calculate the Human Capital Index (HCI)
The data is from September 2020 from....

```python
# In this cell load your data into a Table

my_data_raw = Table().read_table('hci_data_september_2020.csv').select('Region','Income Group', 'Probability of Survival to Age 5', 'Expected Years of School', 'Harmonized Test Scores', 'Adult Survival Rate') # SOLUTION

def generalize_incomeGroup(my_income_group):
    
    if my_income_group == 'Upper middle income':
        my_income_group = 'High income'
    elif my_income_group == 'Lower middle income':
        my_income_group = 'Low income'
    return my_income_group

color_table = Table().with_columns(
    'HIGH/LOW', make_array('High income', 'Low income'),
    'Color', make_array('darkblue', 'gold')
)

income_HIGH_LOW = my_data_raw.apply(generalize_incomeGroup,'Income Group')
my_data_raw = my_data_raw.with_column('HIGH/LOW',income_HIGH_LOW )
my_data_raw = my_data_raw.join('HIGH/LOW', color_table)
my_data_raw

my_data_raw
```

It is very difficult to look at this data in tabular format. Let's look at it by plotting 2 variables and looking at possible associations

```python
# Without looking at the regions, can we see a correlation between the expected years of school and the survival rate
my_data_raw.scatter('Expected Years of School', 'Adult Survival Rate')
my_data_raw.scatter('Expected Years of School', 'Probability of Survival to Age 5')
```

**Question 1.3** When looking at the different plots above, we can see some association between the different variable (probability of survival to 5, adult survival), which variable seems to be affecting the probability of survival the most? Explain in a couple of words.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

Now let's have a look by comparing the data by Income group (High versus Low income).

```python
# my_data_raw.scatter('Expected Years of School', 'Adult Survival Rate', group="Income Group")
my_data_raw.scatter('Expected Years of School', 'Adult Survival Rate', group="HIGH/LOW") 
my_data_raw.scatter('Adult Survival Rate', 'Probability of Survival to Age 5', group="HIGH/LOW")
#my_data_raw.scatter('Expected Years of School', 'Probability of Survival to Age 5', group="Income Group")

ax = plt.figure(figsize=(8,8)).add_subplot(111, projection='3d')
ax.scatter(my_data_raw.column('Expected Years of School'), 
           my_data_raw.column('Adult Survival Rate'), 
           my_data_raw.column('Probability of Survival to Age 5'),
           c = my_data_raw.column('Color'));


```

TODO: Add question here

```python
# Another way of visualizing
my_data_raw.scatter('Expected Years of School', 'Adult Survival Rate', fit_line=True)


```

TODO: Question??

TODO: Add a section explaining correlation coefficient

```python
# Helper function to calculate correlation coefficient
def standard_units(arr):
    return (arr- np.mean(arr)) / np.std(arr)

def correlation(tbl, col1, col2):
    r = np.mean(standard_units(tbl.column(col1)) * standard_units(tbl.column(col2)))
    return r

```

```python
# Calculating the correlation coefficient for 'Expected Years of School' versus 'Adult Survival Rate' 
cor = correlation(my_data_raw, 'Expected Years of School', 'Adult Survival Rate')
cor
```

**Question:** TODO

```python
# Calculating the correlation coefficient for 'Expected Years of School' versus 'Probability of Survival to Age 5' 
cor = ...
cor
```

```python
grader.check("q1_4")
```

TODO: Add another question comparing the correlations?

**Congratulations**: You're DONE with this assignment! Nice work. 
Time to submit.

Tests for q1_4 (points: None)
```
>>> # Correlation coefficient should be between or equal to -1 and 1;
>>> -1 <= cor < 1
True
```
