```python
# Initialize Otter
import otter
grader = otter.Notebook("introVideo.ipynb")
```

<img src="West-Lawn.png" alt="drawing" width="500" height="250"/>

# 1. What is a Jupyter notebook?
This webpage is called a Jupyter notebook. A notebook is a place to write code and view their results, and also to write text.
In a notebook, each rectangle containing text or code is called a *cell*.

**Text cells** (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.
After you edit a text cell, click the "run cell" button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. 

**Code cells** contain code in the Python 3 language. Running a code cell will execute all of the code it contains.
To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶| or hold down `shift` + `return`.

**Activity 1:** This is a text cell. It is the cell type where we can type text that isn't code. Go ahead and double click in this cell and you will see that you can edit it. 

**Type something here:** Hello How are you

**Activity 2:** Click on the code cell below and run the code:

```python
#This cell is a code cell. It is where we can type code that can be executed.
#The hashtag at the start of this line makes it so that this text is a comment not code. 

print("Hello, World! \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

And this one:

```python
# This coding cell imports some python libraries that we will be using throughout this notebook
# Don't worry about what they are, just run this cell before running any other cells below this one
from datascience import *
import numpy as np

%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

import folium
def createMap(locationInit, zoomInit, marker1Location, marker1Name):
    myMap = folium.Map(location=locationInit, tiles="OpenStreetMap", zoom_start= zoomInit)
    folium.Marker(marker1Location, popup = marker1Name).add_to(myMap)
    return myMap

from IPython.display import YouTubeVideo

```

## Air Quality in California

The table California_airquality.csv contains data on air quality in California 1/1/2020 through 9/10/2020 of the air quality index in California as measured by PM2.5, the main pollutant from fire (data is from 2 sources: AQS, AirNow). I found the dataset on Kaggle, direct link here:https://www.kaggle.com/thaddeussegura/california-air-quality-2020-through-sept10th. Kaggle is an open source/free platform where anyone interested in data science can share code and data sets (https://www.kaggle.com/datasets).

```python
CalAQI = Table.read_table('California_airquality.csv')
onlyAQI = CalAQI.select('Date', 'Site Name', 'DAILY_AQI_VALUE')
onlyAQI
```

```python
# look for when Santa Barbara had high AQI during this time
onlySB = onlyAQI.where('Site Name', 'Santa Barbara')
onlySB
```

```python
onlySB.scatter('Date','DAILY_AQI_VALUE')
```

Run the code cell below to watch the video "Why all world maps are wrong"

```python
# Why all world maps are wrong
# The original URL is: 
#https://youtu.be/kIID5FDi2JQ
YouTubeVideo("kIID5FDi2JQ")
```

```python
# Let's create a map
SB_lat = 34.4203
SB_long = -119.6989
createMap([33,-120], 3, [SB_lat, SB_long], 'Santa Barbara')
```

**Question 2.3:** What do geographers call the projection of the 2D map created above?
1. Mercator Projection
2. True Projection
3. Gall-Peters Projection
4. Winkel Tripel Projection

Answer in the code cell below, by setting the name (variable) `projection` equal to 1, 2, 3 or 4.

```python
projection = ...
```

```python
grader.check("q2_3")
```

Tests for q2_3 (points: None)
```
>>> projection == 1
True
```
